# Assignment 2

### Instructions
- Do NOT change the filename.
- Do NOT add or copy cells before the "Sandpit". Only add your own cells after the "Sandpit".
- Enter your name below.

### Enter your name here
Name: 

```python

```

## Task 1: String manipulation
###### (2 points)
- Write code that creates a BHT email address, in this format:
  first 2 letters in first name in lower case, followed by first 2 letters in last name in lower case, followed by "1234"
- Example: "Butter Ball" gets the email address "buba1234@bht-berlin.de"
- (We haven't talked about functions yet, but we need this construct here for auto-grading.)

<!-- grade_id: cell-9e08127856761ea8, points: 0, grade: false, solution: true -->
```python
def create_email_address(first_name, last_name):
    # final line of code should look like this:
    # return a_string_with_email_address
    f_l = first_name[0:2]
    l_l = last_name[0:2]
    new_mail = f_l + l_l + "1234@bht-berlin.de"
    mail = new_mail.lower()
    print(mail)
    return mail
```

<!-- grade_id: cell-70f1705262d021b6, points: 1, grade: true, solution: false -->
```python
assert create_email_address("Butter", "Ball") == "buba1234@bht-berlin.de"
```

Output:
```
buba1234@bht-berlin.de
```

<!-- grade_id: cell-6602798f99dcb2b5, points: 1, grade: true, solution: false -->
```python
assert create_email_address("Rainbow", "Unicorn") == "raun1234@bht-berlin.de"
```

Output:
```
raun1234@bht-berlin.de
```

## Task 2: List manipulation
###### (2 points)
Write code the performs the following:
- Put the following strings/words into a list (your choice of variable name), exactly in this order:

``['the', 'dog',  'quick', 'lazy', 'brown', 'a', 'fox', 'over', 'jumped']``

- Using up to __3__ carefully and cleverly chosen list operations, come up with a new list with (variable name __must__ be "new_list") the same words, but in this new order:

``['the', 'quick', 'brown', 'fox', 'jumped', 'over', 'a', 'lazy', 'dog']``

- No loops allowed! (We haven't talked about them yet.)
- Obviously the following approach gets zero points:

```ipython
# Simply assigning the final answer without list manipulation -- zero points!
list = ['the', 'quick', 'brown', 'fox', 'jumped', 'over', 'a', 'lazy', 'dog']
print(list)
```

<!-- grade_id: cell-34110f1a5e8dfb71, points: 0, grade: false, solution: true -->
```python
start_list = ['the', 'dog',  'quick', 'lazy', 'brown', 'a', 'fox', 'over', 'jumped']
# Slice out "dog", "lazy", "a", "over"
# Start at position 0 and take every second element
a_list = start_list[0:len(start_list):2]
# Slice out "fox", "brown", "quick"
# Start at position -2 (second to last), go left and take every second element
b_list = start_list[-2: :-2]
# concatination 
new_list = a_list + b_list

```

<!-- grade_id: cell-d7728d4672965e61, points: 2, grade: true, solution: false -->
```python
assert new_list == ['the', 'quick', 'brown', 'fox', 'jumped', 'over', 'a', 'lazy', 'dog']
```

# Sandpit begins here
### Add your own cells below

```python

```